
## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.
Instructions

- **Download** this notebook as you would any other ipynb file 
- **Upload** to Google Colab or work locally (if you have that set-up)
- **Delete** `raise NotImplementedError()`

- **Write** your code in the `# YOUR CODE HERE` space


- **Execute** the Test cells that contain assert statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)

- **Save** your notebook when you are finished
- **Download** as a ipynb file (if working in Colab)
- **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)



# Major Neural Network Architectures Challenge
## *Data Science Unit 4 Sprint 3 Challenge*

In this sprint challenge, you'll explore some of the cutting edge of Deep Learning. This week we studied several famous neural network architectures: 
recurrent neural networks (RNNs), long short-term memory (LSTMs), convolutional neural networks (CNNs), and Autoencoders. In this sprint challenge, you will revisit these models. Remember, we are testing your knowledge of these architectures not your ability to fit a model with high accuracy. 

__*Caution:*__  these approaches can be pretty heavy computationally. All problems were designed so that you should be able to achieve results within at most 5-10 minutes of runtime locally, on AWS SageMaker, on Colab or on a comparable environment. If something is running longer, double check your approach!

__*GridSearch:*__ CodeGrade will likely break if it is asked to run a gridsearch for a deep learning model (CodeGrade instances run on a single processor). So while you may choose to run a gridsearch locally to find the optimum hyper-parameter values for your model, please delete (or comment out) the gridsearch code and simply instantiate a model with the optimum parameter values to get the performance that you want out of your model prior to submission. 


## Challenge Objectives
*You should be able to:*
* <a href="#p1">Part 1</a>: Train a LSTM classification model
* <a href="#p2">Part 2</a>: Utilize a pre-trained CNN for object detection
* <a href="#p3">Part 3</a>: Describe a use case for an autoencoder
* <a href="#p4">Part 4</a>: Describe yourself as a Data Science and elucidate your vision of AI

____

# (CodeGrade) Before you submit your notebook you must first

1) Restart your notebook's Kernel

2) Run all cells sequentially, from top to bottom, so that cell numbers are sequential numbers (i.e. 1,2,3,4,5...)
- Easiest way to do this is to click on the **Cell** tab at the top of your notebook and select **Run All** from the drop down menu. 

3) If you have gridsearch code, now is when you either delete it or comment out that code so CodeGrade doesn't run it and crash. 

4) Read the directions in **Part 2** of this notebook for specific instructions on how to prep that section for CodeGrade.

____

<a id="p1"></a>
## Part 1 - LSTMs

Use a LSTM to fit a multi-class classification model on Reuters news articles to distinguish topics of articles. The data is already encoded properly for use in a LSTM model. 

Your Tasks: 
- Use Keras to fit a predictive model, classifying news articles into topics. 
- Name your model as `model`
- Use a `single hidden layer`
- Use `sparse_categorical_crossentropy` as your loss function
- Use `accuracy` as your metric
- Report your overall score and accuracy
- Due to resource concerns on CodeGrade, `set your model's epochs=1`

For reference, the LSTM code we used in class will be useful. 

__*Note:*__  Focus on getting a running model, not on maxing accuracy with extreme data size or epoch numbers. Only revisit and push accuracy if you get everything else done! 

```python
# Import data (don't alter the code in this cell)
from tensorflow.keras.datasets import reuters

# Suppress some warnings from deprecated reuters.load_data
import warnings
warnings.filterwarnings('ignore')

# Load data
(X_train, y_train), (X_test, y_test) = reuters.load_data(num_words=None,
                                                         skip_top=0,
                                                         maxlen=None,
                                                         test_split=0.2,
                                                         seed=723812,
                                                         start_char=1,
                                                         oov_char=2,
                                                         index_from=3)

# Due to limited computational resources on CodeGrade, take the following subsample 
train_size = 1000
X_train = X_train[:train_size]
y_train = y_train[:train_size]
```

```python
# Demo of encoding
word_index = reuters.get_word_index(path="reuters_word_index.json")

print(f"Iran is encoded as {word_index['iran']} in the data")
print(f"London is encoded as {word_index['london']} in the data")
print("Words are encoded as numbers in our dataset.")
```

Output:
```
Iran is encoded as 779 in the data
London is encoded as 544 in the data
Words are encoded as numbers in our dataset.
```

```python
len(word_index.values())
```

Output:
```
30979
```

```python
# Imports (don't alter this code)
from tensorflow.keras.preprocessing import sequence
from tensorflow.keras.models import Sequential
from tensorflow.keras.layers import Dense, Embedding, LSTM

# DO NOT CHANGE THESE VALUES 
# Keras docs say that the + 1 is needed: https://keras.io/api/layers/core_layers/embedding/
MAX_FEATURES = len(word_index.values()) + 1

# maxlen is the length of each sequence (i.e. document length)
MAXLEN = 200
```

```python
type(X_train[0])
```

Output:
```
list
```

<!-- grade_id: cell-471d7f5819bebff6, points: 0, grade: false, solution: true -->
```python
# Pre-process your data by creating sequences 
# Save your transformed data to the same variable name:
# example: X_train = some_transformation(X_train)

# YOUR CODE HERE
'''Pad Sequences'''
X_train = sequence.pad_sequences(X_train, maxlen=MAXLEN)
X_test = sequence.pad_sequences(X_test, maxlen=MAXLEN)
print(X_train.shape)
print(X_test.shape)

```

Output:
```
(1000, 200)
(2246, 200)
```

```python
print(y_train.shape)
print(y_test.shape)
```

Output:
```
(1000,)
(2246,)
```

```python
X_train[0]
```

Output:
```
array([   0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    0,    0,    0,    0,    0,    0,    0,    0,    0,
          0,    0,    1,  248,  409,  166,  265, 1537, 1662,    8,   24,
          4, 1222, 2771,    7,  227,  236,   40,   
```

<!-- grade_id: cell-b46c98c26266363a, points: 1, grade: true, solution: false -->
```python
# Visible tests
assert X_train.shape[1] == MAXLEN, "Your train input sequences are the wrong length. Did you use the sequence import?"
assert X_test.shape[1] == MAXLEN, "Your test input sequences are the wrong length. Did you use the sequence import?"
```

### Create your model

Make sure to follow these instructions (also listed above):
- Name your model as `model`
- Use a `single hidden layer`
- Use `sparse_categorical_crossentropy` as your loss function
- Use `accuracy` as your metric

**Additional considerations**

The number of nodes in your output layer should be equal to the number of **unique** values in the sequences you are training and testing on. For this text, that value is equal to 46.

- Set the number of nodes in your output layer equal to 46

<!-- grade_id: cell-5e7ea9089f827793, points: 0, grade: false, solution: true -->
```python
# Build and compile your model here

# YOUR CODE HERE
model = Sequential()
model.add(Embedding(MAX_FEATURES, 64, input_length=MAXLEN))
model.add(LSTM(128, return_sequences=False))
model.add(Dense(46, activation='softmax'))

model.compile(loss='sparse_categorical_crossentropy',
              optimizer='adam',
              metrics=['accuracy'])
model.summary()
```

Output:
```
Model: "sequential"
_________________________________________________________________
Layer (type)                 Output Shape              Param #   
=================================================================
embedding (Embedding)        (None, 200, 64)           1982720   
_________________________________________________________________
lstm (LSTM)                  (None, 128)               98816     
_________________________________________________________________
dense (Dense)                (None, 46)                5934      
=================================================================
Total params: 2,087,470
Trainable params: 2,087,470
Non-trainable params: 0
_________________________________________________________________
```

<!-- grade_id: cell-54f4676c642d2c94, points: 1, grade: true, solution: false -->
```python
# Visible Test
assert model.get_config()["layers"][1]["class_name"] == "Embedding", "Layer 1 should be an Embedding layer."
```

<!-- grade_id: cell-974465c65fe51083, points: 1, grade: true, solution: false -->
```python
# Hidden Test
```

### Fit your model

Now, fit the model that you built and compiled in the previous cells. Remember to set your `epochs=1`! 

<!-- grade_id: cell-10c20c87933d059c, points: 0, grade: false, solution: true -->
```python
# Fit your model here
# REMEMBER to set epochs=1

# YOUR CODE HERE
model.fit(X_train, y_train,
          batch_size=32,
          epochs=1,
          validation_data=(X_test, y_test))
```

Output:
```
32/32 [==============================] - 7s 126ms/step - loss: 3.4129 - accuracy: 0.2685 - val_loss: 2.4293 - val_accuracy: 0.3664
<tensorflow.python.keras.callbacks.History at 0x7f3d5ebf1f50>
```

<!-- grade_id: cell-277a7dc0b08b9a29, points: 1, grade: true, solution: false -->
```python
# Visible Test 
n_epochs = len(model.history.history["loss"])
assert n_epochs == 1, "Verify that you set epochs to 1."
```

## Sequence Data Question
#### *Describe the `pad_sequences` method used on the training dataset. What does it do? Why do you need it?*

<!-- grade_id: cell-92a7ebc76ad66f05, points: 0, grade: true, solution: true -->
In the neural network, inputs must follow a standard structure. In this case, the inputs must be of a value of MAX_FEATURES. If one of the sequences has less than that many inputs, then "pad_sequences" will add on extra zeros at the beginning of the sequence until there are the correct number of inputs. This way, all the sequences will be the same length.

## RNNs versus LSTMs
#### *What are the primary motivations behind using Long-ShortTerm Memory Cell unit over traditional Recurrent Neural Networks?*

<!-- grade_id: cell-bfe45496d78d39bb, points: 0, grade: true, solution: true -->
A Recurrent Neural Network has a Vanishing Gradient Problem, where the early neurons will not learn as much as the later neurons. This is especially troubling for training data that is sequential, and results in the model having short term memory. LSTM solves this problem by using gates, which help the model learn long term dependencies. 

## RNN / LSTM Use Cases
#### *Name and Describe 3 Use Cases of LSTMs or RNNs and why they are suited to that use case*

<!-- grade_id: cell-eeaef2336d124b88, points: 0, grade: true, solution: true -->
Three example use cases for LSTM or RNN:

Stock price forecasting.

Speech recognition.

Sentiment classification.

These are examples of sequential data, where the order of the sequences matter for the model to train on. LSTM and RNN have the ability to remember previous neurons in its training process, and take the order of the features into account.

<a id="p2"></a>
## Part 2- CNNs

### Find the Frog

Time to play "find the frog!" Use Keras and [ResNet50v2](https://www.tensorflow.org/api_docs/python/tf/keras/applications/resnet_v2) (pre-trained) to detect which of the images with the `frog_images` subdirectory has a frog in it.

<img align="left" src="https://d3i6fh83elv35t.cloudfront.net/newshour/app/uploads/2017/03/GettyImages-654745934-1024x687.jpg" width=400>

The skimage function below will help you read in all the frog images into memory at once. You should use the preprocessing functions that come with ResnetV2, and you should also resize the images using scikit-image.

### Reading in the images

The code in the following cell will download the images to your notebook (either in your local Jupyter notebook or in Google colab).

```python
# Prep to import images (don't alter the code in this cell)
import urllib.request

# Text file of image URLs
text_file = "https://raw.githubusercontent.com/LambdaSchool/data-science-canvas-images/main/unit_4/sprint_challenge_files/frog_image_url.txt"
data = urllib.request.urlopen(text_file)

# Create list of image URLs
url_list = [] 
for line in data:
    url_list.append(line.decode('utf-8'))
```

```python
# Import images (don't alter the code in this cell)

from skimage.io import imread
from skimage.transform import resize 

# instantiate list to hold images
image_list = []

### UNCOMMENT THE FOLLOWING CODE TO LOAD YOUR IMAGES

#loop through URLs and load each image
for url in url_list:
    image_list.append(imread(url))

## UNCOMMENT THE FOLLOWING CODE TO VIEW AN EXAMPLE IMAGE SIZE
#What is an "image"?
print(type(image_list[0]), end="\n\n")

print("Each of the Images is a Different Size")
print(image_list[0].shape)
print(image_list[1].shape)
```

Output:
```
<class 'numpy.ndarray'>

Each of the Images is a Different Size
(2137, 1710, 3)
(3810, 2856, 3)
```

```python
print(image_list[2].shape)
print(image_list[3].shape)
```

Output:
```
(3456, 4608, 3)
(2500, 3335, 3)
```

### Run ResNet50v2

Your goal is to validly run ResNet50v2 on the input images - don't worry about tuning or improving the model. You can print out or view the predictions in any way you see fit. In order to receive credit, you need to have made predictions at some point in the following cells.

*Hint* - ResNet 50v2 doesn't just return "frog". The three labels it has for frogs are: `bullfrog, tree frog, tailed frog`

**Autograded tasks**

* Instantiate your ResNet 50v2 and save to a variable named `resnet_model`

**Other tasks**
* Re-size your images
* Use `resnet_model` to predict if each image contains a frog
* Decode your predictions
* Hint: the lesson on CNNs will have some helpful code

**Stretch goals***
* Check for other things such as fish
* Print out the image with its predicted label
* Wrap everything nicely in well documented functions

## Important note!

To increase the chances that your notebook will run in CodeGrade, when you **submit** your notebook:

* comment out the code where you load the images
* comment out the code where you make the predictions
* comment out any plots or image displays you create

**MAKE SURE YOUR NOTEBOOK RUNS COMPLETELY BEFORE YOU SUBMIT!**

```python
# Imports
import numpy as np
import matplotlib.pyplot as plt

from tensorflow.keras.applications.resnet_v2 import ResNet50V2 # <-- pre-trained model 
from tensorflow.keras.preprocessing import image
from tensorflow.keras.applications.resnet_v2 import preprocess_input, decode_predictions
```

```python
#Show the examples image, then comment it out

# plt.imshow(image_list[2])
```

<!-- grade_id: cell-2715a95f2bcdf3f8, points: 0, grade: false, solution: true -->
```python
# Code from the CNN lecture might come in handy here! 

# YOUR CODE HERE
# def process_img(img):
#   return image.load_img(img, target_size=(2000, 3000))

img = image_list[2]

x = np.resize(img, (224, 224, 3))
x = image.img_to_array(x)
x = np.expand_dims(x, axis=0)
x = preprocess_input(x)
resnet_model = ResNet50V2(weights='imagenet')
features = resnet_model.predict(x)
results = decode_predictions(features, top=4)[0]
print(results)

# def img_pretrain(img):
#   x = np.resize(img, (224, 224, 3))
#   x = image.img_to_array(x)
#   x = np.expand_dims(x, axis=0)
#   x = preprocess_input(x)
#   resnet_model = ResNet50V2(weights='imagenet')
#   features = resnet_model.predict(x)
#   results = decode_predictions(features, top=4)[0]
#   print(results)
```

Output:
```
[('n04590129', 'window_shade', 0.9344607), ('n04041544', 'radio', 0.004158903), ('n03223299', 'doormat', 0.003934468), ('n03721384', 'marimba', 0.0037186353)]
```

<!-- grade_id: cell-6e0982cb9f7775ef, points: 1, grade: true, solution: false -->
```python
# Visible test
assert resnet_model.get_config()["name"] == "resnet50v2", "Did you instantiate the resnet model?"
```

<a id="p3"></a>
## Part 3 - Autoencoders

**Describe a use case for an autoencoder given that an autoencoder tries to predict its own input.**

<!-- grade_id: cell-1ec34a8c8251db51, points: 0, grade: true, solution: true -->
Autoencoders can be applied to problems such as facial recognition. The autoencoder learns an encoded representation of the data of an image, after dimensionality reduction, and then reconstructs the representation to be as similar to the original as possible. In this way, it takes away the noise of the image, and can be used to identify the target in other images.

<a id="p4"></a>
## Part 4 - More...

**Answer the following questions, with a target audience of a fellow Data Scientist:**

- What do you consider your strongest area as a Data Scientist?
- What area of Data Science would you most like to learn more about, and why?
- Where do you think Data Science will be in 5 years?

A few sentences per answer is fine - only elaborate if time allows.

<!-- grade_id: cell-7133ec302afe51d8, points: 0, grade: true, solution: true -->
My strongest area as a Data Scientist problem solving. Knowing the fundamental skills of Data Science, I'm able to achieve most of the goals and questions set in curriculum assignments with a mix of learned knowledge, experimentation, and research.

The area I'm most interested in learning more about is Machine Learning. I would like to learn more about the statistical theories and preprocessing of data required to create strong forecasting models.

In five years, I think Data Science will have grown to become an even larger part of every aspect of our society, even more than it is now. Presently, Data Science might be out of reach for organizations or communities without the resources to utilize it, but as data continues to accumulate, and data scientists become more commonplace, perhaps even small businesses and groups will be able to apply DS to further their causes.

## Congratulations! 

Thank you for your hard work, and [congratulations](https://giphy.com/embed/26xivLqkv86uJzqWk)!!! You've learned a lot, and you should proudly call yourself a Data Scientist.
